# COGS 108 - Discusison Lab Exercise 1: Conda & git / GitHub

This assignment covers basic usage of git / GitHub, and getting a working Conda envirornment on your local computer. 

This assignment is out of 2 points


### Discussion lab exercise rules
Discussion labs are meant to be a place for you to learn and explore. It is ok for you to ask questions of other students and your TAs or seek information beyond the classroom environment. If you use AI assistance, please make sure you understand the solution and can reproduce it on your own. 

### AUTOGRADER WARNINGS

This notebook will be graded automatically, its easy to do things that break the autograder and prevent you from getting the points you deserve.

- DO NOT change the name of this file
- DO NOT copy-paste or delete entire cells
    - You CAN add additional cells, if you want to test things out or need a scratch pad
- Follow directions carefully, typos and mistakes will lose points
- Save regularly, autosave is unreliable
- When you are finished save and on the pull-down menu choose Kernel->Restart Kernel and Run all.  Check to make sure all your code executed without error; with notebooks it's possible to have out of order execution which can lead you to think that your code works when it really doesn't
- Don't forget to submit this notebook on the Assignments tab. The UI is terrible, its easy to press the wrong button and resubmit an old exercise instead of submitting this one. Make sure you see a new submission timestamp on THIS exercise before you log off.

### How to complete assignments

Whenever you see:

```python
# YOUR CODE HERE
raise NotImplementedError()
```

You need to **replace** these lines of code with some code that answers the question correctly. Make sure you remove the 'raise' line or even your correct answer will still lose points!!

The answer needs to be in this cell (the one with `# YOUR CODE HERE`), but you can also add extra cells to explore / investigate things if you need / want to. 

### How to Score Points

Any cell with `assert` statements in it is a test cell. You should not try to change or delete these cells. Note that there might be more than one assert that tests a particular question, and there may be some additional hidden `assert` statements to test your code that you cannot see.

Reading assert errors if and when they fail is a great way to debug your code and find out what went wrong.

Note that some cells, including the test cells, may be read only, which means they won't let you edit them. If you cannot edit a cell - that is normal, and you shouldn't need to edit that cell.

All outside packages/modules that will be used will be specified. You may not use other libraries in the assignments.

Finally, note that questions have points as specified in the detailed instructions. 

## Part I - Installing Python and managing virtual environments using Anaconda ON YOUR LOCAL MACHINE

__NOTE 1__ These instructions are to install on your local laptop/desktop. You're NOT installing on Datahub, where you are right now looking at these instructions :)

__NOTE 2__ Some people struggle with this. Come to section for help.  Also if you're struggling with the install for way too long, please note it is possible to answer the multiple choice questions for points below with just theoretical knowledge if you read the documentation.  Its better to actually do it and understand it first hand, but if its a choice between turning something in on time or not... probably better to turn in on time and fix the conda install later.

__NOTE 3__: Yeah all these instructions say FA25, not WI26.  Its still the same setup, don't worry, it works fine for this quarter.  Don't change names, just follow instructions as they are given 

### why do this (motivation, instructions are later)
It is very helpful to have a local copy of Python instead of always relying on Datahub. That's our main goal here. Now we could just download Python and some other tools and be done, but that's going to cause problems for you later.

Programming professionals don't have just one set of tools.  Sometimes your project depends on features only found in Python 3.9 and later, and sometimes your project relies on an older library that won't work with anything newer than Python 3.7  

If you're going to become a professional who uses programming on many different projects you will quickly realize you can't have just one installation. You need to have one environment for project A, and a different environment for project B, etc.. This is why there are many different solutions for creating and switching between *virtual environments*.

[Anaconda](https://anaconda.org/) is a Python distribution (Python plus a bunch of tools and libraries) that is popular in the Data Science / Machine Learning field. It's built on top of `conda`, an open source virtual environment manager. We are going to learn some conda basics in this section while getting you to install Python on your local system 

When you create a new environment using `conda` you tell it what packages you want in there, e.g. Python, pandas, mysql, etc.. If you don't specify which version number for a package it will use the newest available version that works given the other packages you're choosing (there are sometimes dependencies between packages). You can specify partial version numbers (python=3 or python=3.9) or full version numbers (python=3.9.11).  If you specify only a partial version number `conda` will take the newest usable version that matches the partial number.  You can create new environments directly on the command line (see Q2) or through a YAML file which defines the environment setup.  A YAML file is what we will use below to do...  

### Installation and setup

If you DO NOT have `git` installed on your local machine please do so first.

If you DO NOT have a current Anaconda installation already on your machine then download and install Anaconda on your local machine
https://www.anaconda.com/download

1) Open a terminal window.  Exactly how you do this will be different depending on your operating system. A little searching should show you how.
1) Change directory to someplace in your filesystem that you would find useful to hold COGS 108 related materials
1) type `git clone https://github.com/COGS108/Resources.git` <br> this repo contains a file that describes a conda environment that will mimic the setup we have on Datahub
1) type `cd Resources` so we can use that file to create your new COGS108 environment
1) familiarize yourself with the documentation on how to use conda to create a new python environment https://docs.conda.io/projects/conda/en/latest/user-guide/tasks/manage-environments.html#
1) open the file "conda_env_COGS108_FA25.yml" in a file editor of your choice and take a look at it.  Try to understand how it works to specify exactly what version of each package it is going to install in the environment. If you need help understanding the file, ask questions of your TA or try reading the first part of this https://carpentries-incubator.github.io/introduction-to-conda-for-data-scientists/04-sharing-environments/index.html    
1) OK let's make the new conda environment! In your terminal type `conda env create -f conda_env_COGS108_FA25.yml --solver=libmamba`
1) If the command above produces any errors, ask for help from a TA during the discussion section or come to office hours or post a full bug report to the D1 tag on Ed. <br> If making a bug report include a screenshot or copy/paste of
    - the conda create command,
    - its error message,
    - the outputs of the terminal command `conda --version`
    - the outputs of the terminal command `uname -a`
1) If no errors happened lets see the environments you have available to you.  Type `conda env list` and one of them should be "COGS108_FA25"
2) OK lets try out the new environment.  Type `conda activate COGS108_FA25`
3) Double check that your new environment works by typing `jupyter-lab` and after a few moments you should get a new web browser window with a local copy of Jupyter!
4) We don't need the Jupyter server right now... the only goal is to make sure you have it ready for another day.  After seeing that Jupyter is working, go to the pull down menus, choose File -> Shut Down and close the browser window when it says you're allowed to.
5) Lets leave the COGS 108 environment and go back to the base conda environment.  In the terminal window type `conda activate base`.  In the future if you open a new terminal window it will default to being in the base environment

### Q1 (0.2 points)

Which of the following choices shows you the terminal commands to correctly get into your COGS108 environment and start a Jupyter notebook whenever you need it.  FYI: In a terminal, putting two commands on a single line seperated by a  `;` is the same thing as pressing return in between the two commands 

A) `conda start COGS108_FA25; jupyter-lab`<br>
B) `conda activate; jupyter-lab`<br>
C) `conda activate COGS108_FA25; jupytr`<br>
D) `conda activate COGS108_FA25; jupyter-lab`<br>

Encode your answer in a string variable named `Q1` using the cell immediately below.  Make sure your answer passes the autograder tests in the cell after that!

<!-- grade_id: cell-d6f3e5890e916615, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-19353113982ed038, points: 0.2, grade: true, solution: false -->
```python
# Autograder tests. 
# If this cell executes with no output you passed the tests!
assert Q1, "variable Q1 does not exist"
assert isinstance(Q1,str), "variable Q1 is not a string variable"
assert Q1 in ['A', 'B', 'C', 'D'], "variable Q1 is not one of the allowed multiple choice answers"

```

### Q2 (0.2 points)

Look at the first section of the conda documentation here https://docs.conda.io/projects/conda/en/latest/user-guide/tasks/manage-environments.html#

What one line command would make a new conda environment called "my_project" that has the following setup?

- python = 3.12
- pandas (newest version)
- dask (newset version)

Put your answer in a string called Q2


<!-- grade_id: cell-585ba3f5c9a43ea6, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-9759b9cf63f68234, points: 0.2, grade: true, solution: false -->
```python
# Autograder tests. 
# If this cell executes with no output you passed the tests!

assert Q2, "variable Q2 does not exist"
assert isinstance(Q2,str), "variable Q2 is not a string variable"
assert Q2.startswith('conda'), "variable Q2 is not a conda command"
assert 'create' in Q2, "Variable Q2 doesn't create an environment"

```

## Q3 (0.2 points)

Look at the documentation here https://docs.conda.io/projects/conda/en/latest/user-guide/tasks/manage-pkgs.html#

Lets say you have an existing conda environment named "my_project" that needs the package seaborn, but doesn't have it installed right now. 

Pick the choice below that shows how could you add the newest available seaborn, assuming you are currently in the base environment.

A) `conda activate my_project; conda install seaborn`<br>
B) `conda install seaborn`<br>
C) `conda install -n my_project seaborn`<br>
D) answers A and C are both correct <br>
E) all of the above are correct

Put your answer in a string called Q3

<!-- grade_id: cell-90ef16c52b5435b2, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-2b2c00036ffe597f, points: 0.2, grade: true, solution: false -->
```python
# Autograder tests. 
# If this cell executes with no output you passed the tests!
assert Q3, "variable Q3 does not exist"
assert isinstance(Q3,str), "variable Q3 is not a string variable"
assert Q3 in ['A', 'B', 'C', 'D', 'E'], "variable Q3 is not one of the allowed multiple choice answers"

```

### Q4 (0.2 points)

Lets verify that you've got a good installation of the COGS 108 environment.

First off activate the COGS 108 environment.  Now lets see what version of statsmodels you have installed.  Type `conda list statsmodels` and place the version number in a STRING called Q4.  

<!-- grade_id: cell-a3635af2329cbdab, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-46ef1d49482f7128, points: 0.2, grade: true, solution: false -->
```python
# Autograder tests. 
# If this cell executes with no output you passed the tests!

assert Q4, "variable Q4 does not exist"
assert isinstance(Q4,str), "variable Q4 is not a string variable... you probably made Q4 a float accidentally!"

```

### Q5 (0.2 points)

You now know how to switch between environments (activate). You also know how to list what version of a library you have in the current environment.

So what's the difference between the python version in your base environment and the python version in your COGS 108 environment?

A) They are the same <br>
B) base > COGS 108 version <br>
C) COGS 108 > base version <br>
D) I can't figure out how to answer this

Place your answer as a string in a variable called Q5

<!-- grade_id: cell-fe5cbd5fd80512fd, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-60750d4fdb0716d1, points: 0.2, grade: true, solution: false -->
```python
# Autograder tests. 
# If this cell executes with no output you passed the tests!
assert Q5, "variable Q5 does not exist"
assert isinstance(Q5,str), "variable Q5 is not a string variable"
assert Q5 in ['A', 'B', 'C', 'D'], "variable Q5 is not one of the allowed multiple choice answers"

```

###  A word of caution about using conda

**DANGER** NEVER mess around with your base environment if you can avoid it. Base environment runs conda, you shouldn't use base for programming projects. Don't add new packages or update things randomly in base.

If you screw up your base environment your conda itself may stop working right.  I have done this, it sucks to figure out wtf is even going on, and then it sucks more to delete everything and start over.

From now on you usage pattern for conda should look like this:
- new software project?  make a dedicated new conda environment especially for that project and use it
- new working environments can be made by copying old ones or you can make them from scratch
- cloning somone else's repo or downloading open source software?  if it comes with a conda environment file use it to setup a new conda environment for that software.
- you only ever update base environment to keep conda running and updated (See conda docs on how to do this after you've had conda for a year or so).  Base environment is not for doing work.
   

## Part II - Git and GitHub

In this section, we will get familiar using Github and Git for version control. Git is very helpful for tracking changes in your code files, who made those changes, and collaborating with others for coding. Github is the platform that uses Git to host your files, changes, and collaboration on the internet. 

We will go over some of the most common Github tasks: 
- branch
- add
- commit
- merge
- pull request

### Instructions
1. Create a GitHub username if you don't have already have one and login to GitHub.
2. Complete the __[Github Skills Intro](https://github.com/skills/introduction-to-github)__ <--- RIGHT CLICK and open this repo in a new browser tab!!!
3. In the new browser tab there are instructions on how to do the exercise. Follow the instructions EXACTLY and ALL THE WAY TO THE END of the exercise or you will not get your points.
4. You are done with the exercise when the repo README file tells you "Congratulations, you've completed this course and joined the world of developers!"

**Some Important Notes** 
- The first thing the repo's instructions say to do is to right click that green "Start course" button... clicking the button makes you a copy of the tutorial repo. While you're making this copy for yourself leave all the default setup alone, just click the button that says "Create repository".  Particularly:
    - **DO NOT CHANGE REPO NAME** when you do this. Leave it as the default: "skills-introduction-to-github".  if you change the name the autograder will not work
    - **DO NOT CHANGE THE REPO TO PRIVATE** we can only grade public repos
- The instructions suggest you should open two windows on the Skills Intro... one to keep the instructions visible to you and another to actually execute the steps.  This is very important... once you try to do something (commit, pull request, etc) you will be on a new webpage and the instructions disappear.  It's hard to keep all the steps in your head, so it's crticial to have a 2nd window open with the instructions for when you need to double-check things
- Use the **exact file names**, **branch names**, and **commit messages** provided in the instructions; deviations can cause your progress to fail the hidden tests and result in point deductions.

### Q6 (1 point)

Once you've completed the exercise, please provide us your Github username to validate the completion of the Github Skills Introduction. In the cell below, provide your Github username (not the email address you used to make your account) as a string in the variable ```github_username```. 

<!-- grade_id: cell-3e5eaf5783d7da2b, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-fdf42ac0898e4c4b, points: 0.5, grade: true, solution: false -->
```python
# Autograder tests
# If this cell executes with no output, you passed the tests!
assert github_username, 'Github username is not entered'
assert type(github_username) == str, 'Username is not a string variable'

import requests
url = 'https://github.com/' + github_username
response = requests.get(url)
assert response.status_code == 200, "The github username you gave us does not exist on github"


```

<!-- grade_id: cell-8d0eb9ca91be5462, points: 0.5, grade: true, solution: false -->
```python
# This cell contains hidden autograder tests
# do not delete, move, or edit it!

```

## End of Discussion Lab Exercise #1

When you think you're done, verify that everything is working correctly by doing the following:
1. Save this notebook
2. Choose the pull down menu item Kernel -> Restart and Run All Cells
3. Check that you see the print statement in the cell below. If you see that happy sunglasses face emoji, it means that you have passed all the publicly visible tests. Yay!! 

```python
print("\N{SMILING FACE WITH SUNGLASSES}")
```